---
title: 'Chapter 1'
description: 'Chapter description goes here.'
---

## Example coding exercise

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

This is an example exercise.

`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

---

## Introduction

```yaml
type: VideoExercise
key: 85c6832af0
xp: 50
```

`@projector_key`
805f9d203e6629132e5b1465b568c395
