---
title: 'Chapter 2'
description: ""
---

## Insert exercise title here

```yaml
type: MultipleChoiceExercise
key: af361b96fb
xp: 50
```

what is the right answer between these

`@possible_answers`
- this is the first 
- this is the second
- this is the third

`@hint`


`@pre_exercise_code`
```{python}

```

`@sct`
```{python}

```
